```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 01: Monoalphabetic Ciphers and Cryptanalysis

## Overview
This homework will cover the topics taught in Lessons 01 - 12. It is a mix of coding exercises to be completed in this Jupyter notebook, and free-response questions to be completed on a corresponding handout.

## Allowable resources
You may use any programming techniques covered in class to answer the following questions on this programming portion of this assignment, however, pay careful attention to the directions for each question to ensure you are using the specific approach requested, when applicable. You are only allowed to use your notes and refer to previous lesson materials and activities as resources to complete this assignment. No artificial intelligence tools, textbooks other than the book for this course, or internet resources may be used to complete this programming portion of the assignment.

## Part 1: Perfect Numbers

A positive integer $n$ is called a **perfect number** if all divisors of $n$ which are less than $n$ sum to equal $n$. 

Put another way, a perfect number is a positive integer that is equal to the sum of its positive factors, excluding the number itself. For example, $6$ is a perfect number because the factors of $6$ that are less than $6$ are $1$, $2$, and $3$, which sum to $6$ ($6 = 1 + 2 + 3$).

### Question 1.1

Write a function named `perfect` which accepts an integer `n` as an input. The function `perfect` should return `True` if the integer `n` is a perfect number, and return `False` if integer `n` is not a perfect number.

For example:

```python
>>> perfect(6)
True

>>> perfect(28)
True

>>> perfect(15)
False
```

```python
def perfect(n):
    ...
```

```python
grader.check("q1_1")
```

### Question 1.2
Write a function `all_perfects_less_than` which accepts an integer `k` and returns a **list** of all perfect numbers less than `k`.

For example:
```python
>>> all_perfects_less_than(50)
[6, 28]

>>> all_perfects_less_than(500)
[6, 28, 496]

>>> all_perfects_less_than(5)
[]
```

```python
def all_perfects_less_than(k):
    ...
```

```python
grader.check("q1_2")
```

## Part 2: Bar Charts and Chi Squared scoring

In this part, you'll be using chi squared scoring to perform some cryptanalysis. Run the code cell below to load in the following functions we've previously written in class:
* `caesar`
* `chi_squared`

**Note:** It's important that you use the provided functions, not functions you've written in class, to ensure that your results are consistent with the automatic grader checks.

```python
def caesar(text, key, decrypt=False, LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ'):
    
    final_text = ''

    for char in text:
        text_numerical = LETTERS.find( char )
        
        if decrypt:
            new_text_numerical = (text_numerical - key) % len(LETTERS)
        else:
            new_text_numerical = (text_numerical + key) % len(LETTERS)
        
        final_text += LETTERS[new_text_numerical]

    if decrypt:
        return final_text
    else:
        return final_text
```

```python
def chi_squared(text, LETTERS='ABCDEFGHIJKLMNOPQRSTUVWXYZ'):

    letter_percentage =[0.082, 0.015, 0.028, 0.043, 0.127, 0.022, 0.02, 0.061, 0.07, 0.0015, 0.0077, 0.04, 0.024, 0.067, 0.075, 0.019, 0.00095, 0.06, 0.063, 0.091, 0.028, 0.0098, 0.024, 0.0015, 0.02, 0.00074]
    total = 0
    
    for char in LETTERS:
        actual = text.count(char)
        expected = letter_percentage[LETTERS.find(char)] * len(text)
        partial_score = (expected - actual)**2 / expected
        total += partial_score
        
    return total
```

<!-- BEGIN QUESTION -->

### Question 2.1

You intercepted the following ciphertext that you suspect was encrypted using the Caesar cipher. Notice that this ciphertext is already cleaned for analysis.

```python
ciphertext = 'LZWVWTSLWLGMJFSEWFLLZWTMKVWHSJLWVSLXANWLZAJLQSEUSJJQAFYGMJVWTSLWLWSELZJGMYZLZWVSJCFWKKLGOSJVLZWKLSLWUZSEHAGFKZAHWAYZLTDWSJQWQWVKLMVWFLKUDMLUZWVLJSNWDEMYKGXUGXXWWSFVJWNAWOWVFGLWUSJVKTQLZWVAEGNWJZWSVDAYZLKJWEWETWJALKFGLBMKLSTGMLXSULKUGSUZDSOKGFKSAVXJGELZWXJGFLKWSLALKSTGMLUGFFWULAFYOALZLZWBMVYWKESCWLZWEUSJWSTGMLQGMJSJYMEWFLAFGVVWVLZGMYZEQKLGESUZUZMJFWVOALZSFPAWLQLZAKOSKEQXAJKLLAEWUGEHWLAFYSLLZWKLSLWDWNWDLZWHJWNAGMKFAYZLAVHJSULAUWVEQGHWFAFYKLSLWEWFLMFLADEQJGGEESLWLZJWSLWFWVLGKEGLZWJEWOALZSHADDGOAXAVAVFLDWLZWJKDWWHLZWUGEHWLALAGFOSKZWDVSLUWFLJSDMFANWJKALQALKAEHGKAFYTJAUCTMADVAFYKESCAFYEWXWWDKESDDWJLZSFMKMSDLWSEKXJGESUJGKKLZWKLSLWZMVVDWVAFZSDDOSQKGJUDSAEWVLSTDWKAFLZWKLMVWFLMFAGFSTMRRAFYZANWGXFWJNGMKWFWJYQSFVDSKLEAFMLWHJWHSJSLAGFKGMJXAJKLJGMFVLZWLGHAUOSKUDAESLWHGDAUQOWFLKMJHJAKAFYDQOWDDEQHSJLFWJVSEAWFSFVASVNSFUWVLGLZWIMSJLWJXAFSDKLZWFKGEWZGOLGLZWKWEAXAFSDKOALZWSUZJGMFVEQUGFXAVWFUWYJWOLZWZGMJKGXJWKWSJUZLZWHJSULAUWVWTSLWKLZWXWWVTSUCKWKKAGFKALOSKSDDUGEAFYLGYWLZWJLZWXAFSDJGMFVOSKSFAFLWDDWULMSDWPLJSNSYSFRSSYSAFKLLZWVWXWFVAFYUZSEHAGFKXJGEOWKLDSCWSUSVWEQLZWSMVAWFUWZSVYJGOFLGFWSJDQSZMFVJWVHWGHDWHSJWFLKUGSUZWKWDAEAFSLWVUGEHWLALGJKWNWFSXWOMFANWJKALQKLMVWFLKUMJAGMKSTGMLLZWUGEEGLAGFOZWFLZWBMVYWKSFFGMFUWVVSEAWFSFVEWSKLZWOAFFWJKAUGMDVFLHJGUWKKALSLXAJKLUGSUZDSOKGFKWFLZMKASKLAUWETJSUWXAFSDDQESVWALJWSDOWZSVVGFWALGMJKESDDHMTDAUKUZGGDZSVVWXWSLWVHJANSLWSUSVWEAWKOALZVWVAUSLWVVWTSLWUGSUZWKSFVDSNAKZJWKGMJUWKGFLZWJAVWZGEWLJGHZQFWKLDWVKSXWDQAFEQDSHAJWSDARWVKGEWLZAFYAEHGJLSFLTWAFYMFVWJWKLAESLWVUSFTWSHGOWJXMDSVNSFLSYWFGGFWWPHWULWVMKLGOAFWPUWHLUGSUZDSOKGFOZGZSVTWDAWNWVAFMKXJGELZWTWYAFFAFYKGEWLAEWKLZSLKSDDLZWSVNSFLSYWQGMFWWV'
```

In the code cell below, generate a bar chart of the character frequency in the ciphertext. The height of the bar should indicate the *proportion* of the ciphertext that contains the letter associated with each bar (not the raw count).

```python
# Add your code to this cell so that when it is run, a bar chart appears below
```

<!-- END QUESTION -->

## Question 2.2

In the code cell below, write a block of code that:
1. uses every possible key for the Caesar cipher to decrypt the text,
2. score each decryption using the chi squared method, and
3. store each score to a list named `chi_squared_scores`, in order. 

For example, the first element in the `chi_squared_scores` list, `chi_squared_scores[0]`, should contain the chi squared score for the decryption obtained using a key of 0, `chi_squared_scores[1]` should contain the score when using a key of 1, and so on.

```python
chi_squared_scores = []


print(chi_squared_scores)
```

```python
grader.check("q2_2")
```

### Question 2.3

Based on the list of chi squared scores you computed in the previous question, determine which key was most likely used to encrypt the message. Assign the value of the key to the variable `likely_key` in the code cell below.

```python
likely_key = ...
```

```python
grader.check("q2_3")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: 4)
```
>>> perfect(6) == True
True
```
failure_message: 6 should be perfect
```
>>> perfect(28) == True
True
```
failure_message: 28 should be perfect
```
>>> perfect(496) == True
True
```
failure_message: 496 should be perfect
```
>>> perfect(2) == False
True
```
failure_message: 2 is not perfect
```
>>> perfect(15) == False
True
```
failure_message: 15 is not perfect
```
>>> perfect(1) == False
True
```
failure_message: 1 is not perfect
```
>>> perfect(0) == False
True
```
failure_message: 0 is not perfect
```
>>> perfect(-6) == False
True
```
failure_message: Negative numbers are not perfect


Tests for q1_2 (points: 4)
```
>>> all_perfects_less_than(1) == []
True
```
failure_message: For k < 6, the function should return an empty list
```
>>> all_perfects_less_than(0) == []
True
```
failure_message: For k < 6, the function should return an empty list
```
>>> all_perfects_less_than(2) == []
True
```
failure_message: For k < 6, the function should return an empty list
```
>>> all_perfects_less_than(7) == [6]
True
```
failure_message: For 6 < k <= 28, the only perfect number is 6
```
>>> all_perfects_less_than(28) == [6]
True
```
failure_message: For 6 < k <= 28, the only perfect number is 6
```
>>> all_perfects_less_than(50) == [6, 28]
True
```
failure_message: for k = 50, both 6 and 28 are perfect numbers
```
>>> all_perfects_less_than(500) == [6, 28, 496]
True
```
failure_message: for k = 50, 6, 28, and 496 are perfect numbers
```
>>> all_perfects_less_than(-10) == []
True
```
failure_message: Negative values for k should also return an empty list


Tests for q2_2 (points: 4)
```
>>> len(chi_squared_scores) == 26
True
```
failure_message: chi_squared_scores should have 26 elements
success_message: chi_squared_scores has 26 elements
```
>>> import numpy as np
>>> bool(np.isclose(chi_squared_scores[0], 11476.590345085184))
True
```
failure_message: The first element of chi_squared_scores is incorrect
success_message: The first element of chi_squared_scores is correct


Tests for q2_3 (points: 1)
```
>>> type(likely_key) == int
True
```
failure_message: likely_key is not an integer, but it should be
success_message: likely_key has been assigned to an integer value
